## Week 6 Assignment - W200 Python for Data Science, UC Berkeley MIDS

Write code in this jupyter Notebook to solve each of the following problems. Each problem should have its solution in a separate cell. Please upload this **Notebook**, your **scrabble.py** file, the **sowpods.txt** file and your **score_word** module with your solutions to your GitHub repository in your SUBMISSIONS/week_06 folder by 11:59PM PST the night before class.

## Objectives:

- Read and understand PEP 8 standards
- Use all of your previously gained knowledge together on a single program
- Demonstrate how to import a user made module and function into python from another .py file
- Demonstrate how to input command line arguments into a .py file

## 6-1. PEP 8 Style Guide (reading and response)

Your first task for this week is to write a **250** word reading response to the article below. In addition, please list **3 questions** that you have from the article. Please write your response in a markdown cell, below this cell in the notebook.

The writing response is a free response, so you may write about your reactions. An interesting thing that you saw in the article, something that really stuck out to you, etc.

**Article**: [The PEP 8 Style Guide](https://www.python.org/dev/peps/pep-0008). This document is really important for Python coders because it describes best practices and customs for how one should write Python code. Please read it all and prepare your questions for class.

<!-- grade_id: cell-eaae902b022d20da, points: 20, grade: true, solution: true -->
#### Question 1: why limit all lines to a maximum of 79 characters? Why is number 79?
#### Question 2: what is the major difference between Python 2 and Python 3?
#### Question 3: what does the this mean: Constants are usually defined on a module level?

##### It's interesting to know that no attribute is really private in Python and to learn more about public and internal interface.


## 6-2. Cheating at Scrabble

Write a Python script that takes a Scrabble rack as a command-line argument and prints all valid Scrabble words that can be constructed from that rack, along with their Scrabble scores, sorted by score. Valid Scrabble words are provided in the data source below. A Scrabble rack is made up of a maximum of any 7 characters.

Below are the requirements for the program:
- This needs to be able to be run as a command line tool as shown below (not an input statement!)
- Please name the python file: `scrabble.py`
- Allow anywhere from 2-7 character tiles to be inputted 
- Do not worry about the number of the same tiles (e.g. a user is allowed to input ZZZZZQQ)
- Output the **total** list of words as (score, word) tuples, sorted by the score as shown below
- Output at the end the 'Total number of words:' that can be made with the letters
- Please include a function called `score_word` in a separate module. Import this function into your main solution code.
- You need to handle input errors from the user and suggest what that error might be caused by and how to fix it (i.e. a helpful error message)
- Implement wildcards as either `*` or `?`. That is, let the user specify a wildcard character that can take any value. There can be a total of two wild cards in any user input (one of each character). Only use the `*` and `?` as wildcard characters
- Wildcard characters are scored as 0 points, just like in the real Scrabble game
- Your program should take less than a minute to run with 2 wildcards in the input - if it is more time than that your algorithm is not optimized very well!
- Write docstrings for the functions and puts comments in your code.

Extra Credit (+10 points):
Allow a user to specify that a certain letter has to be at a certain location. Your program must work without it so this is completely optional. For the extra credit, locations of certain letters must be specified at the command line, it may not be some sort of user prompt.  (Please put a sample of how to run your extra credit & comments in the extra credit cell of this notebook below!)

An example invocation and output:
```
$ python scrabble.py ZAEFIEE
(17, feeze)
(17, feaze)
(16, faze)
(15, fiz)
(15, fez)
(12, zee)
(12, zea)
(11, za)
(6, fie)
(6, fee)
(6, fae)
(5, if)
(5, fe)
(5, fa)
(5, ef)
(2, ee)
(2, ea)
(2, ai)
(2, ae)
Total number of words: 19
```

The Data
http://courses.cms.caltech.edu/cs11/material/advjava/lab1/sowpods.zip contains all words in the official SOWPODS word list, one word per line. You should download the word file and keep it in your repository so that the program is standalone (instead of accessing it over the web from Python).

You can read data from a text file with the following code:

```
with open("sowpods.txt","r") as infile:
    raw_input = infile.readlines()
    data = [datum.strip('\n') for datum in raw_input]
```
This will show the first 6 words:
```
print(data[0:6])
```

Please use the dictionary below containing the letters and their Scrabble values:

```
scores = {"a": 1, "c": 3, "b": 3, "e": 1, "d": 2, "g": 2,
         "f": 4, "i": 1, "h": 4, "k": 5, "j": 8, "m": 3,
         "l": 1, "o": 1, "n": 1, "q": 10, "p": 3, "s": 1,
         "r": 1, "u": 1, "t": 1, "w": 4, "v": 4, "y": 4,
         "x": 8, "z": 10}
```

Tips:
- We recommend that you work on this and try to break down the problems into steps on your own before writing any code. Once you've scoped generally what you want to do, start writing some code and if you get stuck, take a step back and go back to thinking about the problem rather than trying to fix lots of errors at the code level. You should only use the Python standard library in this assignment, however any tool in the standard library is fair game.
- If you keep getting stuck check out: https://openhatch.org/wiki/Scrabble_challenge. This is where we got the idea for this assignment and it provides some helpful tips for guiding you along the way. If that link doesn't work you can use the Google cached version here. However, we would recommend that you try to implement this first before looking at the hints on the website.

Good luck!

```python
with open("sowpods.txt","r") as infile:
    raw_input = infile.readlines()
    data = [datum.strip('\n') for datum in raw_input]


print(data[0:6])


scores = {"a": 1, "c": 3, "b": 3, "e": 1, "d": 2, "g": 2,
         "f": 4, "i": 1, "h": 4, "k": 5, "j": 8, "m": 3,
         "l": 1, "o": 1, "n": 1, "q": 10, "p": 3, "s": 1,
         "r": 1, "u": 1, "t": 1, "w": 4, "v": 4, "y": 4,
         "x": 8, "z": 10}

```

Output:
```
['AA', 'AAH', 'AAHED', 'AAHING', 'AAHS', 'AAL']
```

### The code below will test your command line implementation of the scrabble.py code. We've made some of these tests are available for you to try!

```python
# Code for the testing

import subprocess
from nose.tools import assert_equal 
from nose.tools import assert_true
from nose.tools import assert_greater
from nose.tools import assert_less
```

<!-- grade_id: cell-f55d7e98eba40a10, points: 20, grade: true, solution: false -->
```python
""" Checks that the code runs and checks one user error messages 
    (this is just one of many user errors you should be checking for!)
"""

!python scrabble.py  # no rack error


```

Output:
```
Please provide a Scrabble rack between 2 and 7 characters
```

<!-- grade_id: cell-1893c284e867334a, points: 10, grade: true, solution: false -->
```python
""" Does not fail due to trivial mistakes and takes correct wildcard characters """

!python scrabble.py PENguin    # does not fail due to case
!python scrabble.py PEN*?in    # takes wildcards
```

Output:
```
(10, 'penguin')
(9, 'pening')
(8, 'unpeg')
(8, 'genip')
(7, 'unpin')
(7, 'unpen')
(7, 'pung')
(7, 'ping')
(7, 'penni')
(7, 'ingenu')
(6, 'pug')
(6, 'pine')
(6, 'pig')
(6, 'peni')
(6, 'pein')
(6, 'peg')
(6, 'gup')
(6, 'gip')
(5, 'pun')
(5, 'piu')
(5, 'pin')
(5, 'pie')
(5, 'pen')
(5, 'nip')
(5, 'nep')
(5, 'ginn')
(5, 'gien')
(5, 'genu')
(5, 'ennui')
(4, 'up')
(4, 'pi')
(4, 'pe')
(4, 'nine')
(4, 'neg')
(4, 'gun')
(4, 'gue')
(4, 'gnu')
(4, 'gin')
(4, 'gie')
(4, 'gen')
(4, 'eng')
(3, 'uni')
(3, 'ug')
(3, 'nun')
(3, 'nie')
(3, 'inn')
(3, 'gu')
(3, 'gi')
(2, 'un')
(2, 'nu')
(2, 'ne')
(2, 'in')
(2, 'en')
(21, 'ziz')
(19, 'zex')
(19, 'jiz')
(18, 'zinke')
(18, 'phiz')
(18, 'phenix')
(18, 'jinx')
(17, 'wizen')
(17, 'winze')
(17, 'pyxie')
(17, 'pized')
(17, 'hizen')
(17, 'benzin')
(16, 'zineb')
(16, 'zek')
(16, 'unzip')
(16, 'prize')
(16, 'pizes')
(16, 'pixy')
(16, 'pique')
(16, 'piezo')
(16, 'peize')
(16, 'mizen')
(16, 'equip')
(15, 'zips')
(15, 'zinc')
(15, 'zeps')
(15, 'wiz')
(15, 'vixen')
(15,
```

<!-- grade_id: cell-7f1c7fe54a38a19e, points: 5, grade: true, solution: false -->
```python
""" The code should produce a list of all words from the rack with scores """

#test the code in the command line
!python scrabble.py PENGUIN

```

Output:
```
penguin
(10, 'penguin')
(9, 'pening')
(8, 'unpeg')
(8, 'genip')
(7, 'unpin')
(7, 'unpen')
(7, 'pung')
(7, 'ping')
(7, 'penni')
(7, 'ingenu')
(6, 'pug')
(6, 'pine')
(6, 'pig')
(6, 'peni')
(6, 'pein')
(6, 'peg')
(6, 'gup')
(6, 'gip')
(5, 'pun')
(5, 'piu')
(5, 'pin')
(5, 'pie')
(5, 'pen')
(5, 'nip')
(5, 'nep')
(5, 'ginn')
(5, 'gien')
(5, 'genu')
(5, 'ennui')
(4, 'up')
(4, 'pi')
(4, 'pe')
(4, 'nine')
(4, 'neg')
(4, 'gun')
(4, 'gue')
(4, 'gnu')
(4, 'gin')
(4, 'gie')
(4, 'gen')
(4, 'eng')
(3, 'uni')
(3, 'ug')
(3, 'nun')
(3, 'nie')
(3, 'inn')
(3, 'gu')
(3, 'gi')
(2, 'un')
(2, 'nu')
(2, 'ne')
(2, 'in')
(2, 'en')
```

<!-- grade_id: cell-bb7a872eda80e6a5, points: 5, grade: true, solution: false -->
```python
# Autograding test
```

<!-- grade_id: cell-8d2de0735f148301, points: 5, grade: true, solution: false -->
```python
""" The code should produce a list of all words from the rack with scores """
!python scrabble.py PENGU*?

```

```python
!python scrabble.py PENGU*?
```

```python
!python scrabble.py PENGU*
```

Output:
```
(16, 'expugn')
(15, 'quep')
(14, 'zep')
(13, 'jupe')
(11, 'jug')
(11, 'gju')
(10, 'punk')
(10, 'puke')
(10, 'pugh')
(10, 'penk')
(10, 'pegh')
(10, 'jun')
(10, 'jeu')
(9, 'xu')
(9, 'upgone')
(9, 'unpegs')
(9, 'spunge')
(9, 'repugn')
(9, 'pyne')
(9, 'puny')
(9, 'pungle')
(9, 'punce')
(9, 'pudge')
(9, 'plunge')
(9, 'kep')
(9, 'gyp')
(9, 'gunk')
(9, 'gump')
(9, 'ex')
(8, 'yup')
(8, 'yep')
(8, 'upend')
(8, 'unpeg')
(8, 'speug')
(8, 'pye')
(8, 'puy')
(8, 'purge')
(8, 'pungs')
(8, 'punga')
(8, 'puh')
(8, 'puce')
(8, 'pube')
(8, 'pew')
(8, 'pengo')
(8, 'pendu')
(8, 'peh')
(8, 'nuke')
(8, 'neuk')
(8, 'kune')
(8, 'keg')
(8, 'hup')
(8, 'hung')
(8, 'huge')
(8, 'hep')
(8, 'getup')
(8, 'genip')
(8, 'fung')
(8, 'eugh')
(8, 'begun')
(7, 'yug')
(7, 'vug')
(7, 'veg')
(7, 'upgo')
(7, 'unpen')
(7, 'ump')
(7, 'uke')
(7, 'ugh')
(7, 'spug')
(7, 'pup')
(7, 'pung')
(7, 'pugs')
(7, 'pub')
(7, 'prune')
(7, 'pong')
(7, 'plug')
(7, 'ping')
(7, 'pep')
(7, 'pend')
(7, 'pegs')
(7, 'pec')
(7, 'peag')
(7, 'pang')
(7, '
```

```python
!python scrabble.py PENGU*
```

```python
!python scrabble.py PENGU*
```

Output:
```
(16, 'EXPUGN')
(15, 'QUEP')
(14, 'ZEP')
(13, 'JUPE')
(11, 'JUG')
(11, 'GJU')
(10, 'PUNK')
(10, 'PUKE')
(10, 'PUGH')
(10, 'PENK')
(10, 'PEGH')
(10, 'JUN')
(10, 'JEU')
(9, 'XU')
(9, 'UPGONE')
(9, 'UNPEGS')
(9, 'SPUNGE')
(9, 'REPUGN')
(9, 'PYNE')
(9, 'PUNY')
(9, 'PUNGLE')
(9, 'PUNCE')
(9, 'PUDGE')
(9, 'PLUNGE')
(9, 'KEP')
(9, 'GYP')
(9, 'GUNK')
(9, 'GUMP')
(9, 'EX')
(8, 'YUP')
(8, 'YEP')
(8, 'UPEND')
(8, 'UNPEG')
(8, 'SPEUG')
(8, 'PYE')
(8, 'PUY')
(8, 'PURGE')
(8, 'PUNGS')
(8, 'PUNGA')
(8, 'PUH')
(8, 'PUCE')
(8, 'PUBE')
(8, 'PEW')
(8, 'PENGO')
(8, 'PENDU')
(8, 'PEH')
(8, 'NUKE')
(8, 'NEUK')
(8, 'KUNE')
(8, 'KEG')
(8, 'HUP')
(8, 'HUNG')
(8, 'HUGE')
(8, 'HEP')
(8, 'GETUP')
(8, 'GENIP')
(8, 'FUNG')
(8, 'EUGH')
(8, 'BEGUN')
(7, 'YUG')
(7, 'VUG')
(7, 'VEG')
(7, 'UPGO')
(7, 'UNPEN')
(7, 'UMP')
(7, 'UKE')
(7, 'UGH')
(7, 'SPUG')
(7, 'PUP')
(7, 'PUNG')
(7, 'PUGS')
(7, 'PUB')
(7, 'PRUNE')
(7, 'PONG')
(7, 'PLUG')
(7, 'PING')
(7, 'PEP')
(7, 'PEND')
(7, 'PEGS')
(7, 'PEC')
(7, 'PEAG')
(7, 'PANG')
(7, '
```

```python
!python scrabble.py PENGU*
```

Output:
```
(16, 'EXPUGN')
(15, 'QUEP')
(14, 'ZEP')
(13, 'JUPE')
(11, 'JUG')
(11, 'GJU')
(10, 'PUNK')
(10, 'PUKE')
(10, 'PUGH')
(10, 'PENK')
(10, 'PEGH')
(10, 'JUN')
(10, 'JEU')
(9, 'XU')
(9, 'UPGONE')
(9, 'UNPEGS')
(9, 'SPUNGE')
(9, 'REPUGN')
(9, 'PYNE')
(9, 'PUNY')
(9, 'PUNGLE')
(9, 'PUNCE')
(9, 'PUDGE')
(9, 'PLUNGE')
(9, 'KEP')
(9, 'GYP')
(9, 'GUNK')
(9, 'GUMP')
(9, 'EX')
(8, 'YUP')
(8, 'YEP')
(8, 'UPEND')
(8, 'UNPEG')
(8, 'SPEUG')
(8, 'PYE')
(8, 'PUY')
(8, 'PURGE')
(8, 'PUNGS')
(8, 'PUNGA')
(8, 'PUH')
(8, 'PUCE')
(8, 'PUBE')
(8, 'PEW')
(8, 'PENGO')
(8, 'PENDU')
(8, 'PEH')
(8, 'NUKE')
(8, 'NEUK')
(8, 'KUNE')
(8, 'KEG')
(8, 'HUP')
(8, 'HUNG')
(8, 'HUGE')
(8, 'HEP')
(8, 'GETUP')
(8, 'GENIP')
(8, 'FUNG')
(8, 'EUGH')
(8, 'BEGUN')
(7, 'YUG')
(7, 'VUG')
(7, 'VEG')
(7, 'UPGO')
(7, 'UNPEN')
(7, 'UMP')
(7, 'UKE')
(7, 'UGH')
(7, 'SPUG')
(7, 'PUP')
(7, 'PUNG')
(7, 'PUGS')
(7, 'PUB')
(7, 'PRUNE')
(7, 'PONG')
(7, 'PLUG')
(7, 'PING')
(7, 'PEP')
(7, 'PEND')
(7, 'PEGS')
(7, 'PEC')
(7, 'PEAG')
(7, 'PANG')
(7, '
```

```python
!python scrabble.py PENGU*
```

```python
!python scrabble.py PENGU*
```

Output:
```
2 EN
2 NE
2 NU
2 UN
3 GU
3 UG
4 ENG
4 GEN
4 GNU
4 GUE
4 GUN
4 NEG
4 PE
4 UP
5 GENU
5 NEP
5 PEN
5 PUN
6 GUP
6 PEG
6 PUG
7 PUNG
8 UNPEG
```

```python
!python scrabble.py PENGU*
```

Output:
```
(16, 'expugn')
(15, 'quep')
(14, 'zep')
(13, 'jupe')
(11, 'jug')
(11, 'gju')
(10, 'punk')
(10, 'puke')
(10, 'pugh')
(10, 'penk')
(10, 'pegh')
(10, 'jun')
(10, 'jeu')
(9, 'xu')
(9, 'upgone')
(9, 'unpegs')
(9, 'spunge')
(9, 'repugn')
(9, 'pyne')
(9, 'puny')
(9, 'pungle')
(9, 'punce')
(9, 'pudge')
(9, 'plunge')
(9, 'kep')
(9, 'gyp')
(9, 'gunk')
(9, 'gump')
(9, 'ex')
(8, 'yup')
(8, 'yep')
(8, 'upend')
(8, 'unpeg')
(8, 'speug')
(8, 'pye')
(8, 'puy')
(8, 'purge')
(8, 'pungs')
(8, 'punga')
(8, 'puh')
(8, 'puce')
(8, 'pube')
(8, 'pew')
(8, 'pengo')
(8, 'pendu')
(8, 'peh')
(8, 'nuke')
(8, 'neuk')
(8, 'kune')
(8, 'keg')
(8, 'hup')
(8, 'hung')
(8, 'huge')
(8, 'hep')
(8, 'getup')
(8, 'genip')
(8, 'fung')
(8, 'eugh')
(8, 'begun')
(7, 'yug')
(7, 'vug')
(7, 'veg')
(7, 'upgo')
(7, 'unpen')
(7, 'ump')
(7, 'uke')
(7, 'ugh')
(7, 'spug')
(7, 'pup')
(7, 'pung')
(7, 'pugs')
(7, 'pub')
(7, 'prune')
(7, 'pong')
(7, 'plug')
(7, 'ping')
(7, 'pep')
(7, 'pend')
(7, 'pegs')
(7, 'pec')
(7, 'peag')
(7, 'pang')
(7, '
```

<!-- grade_id: cell-16d4589caf032f51, points: 5, grade: true, solution: false -->
```python
# Autograding test
```

<!-- grade_id: cell-ceab3d4433043319, points: 10, grade: true, solution: false -->
```python
""" The code should run in seconds to a few minutes """
import time
start=time.time()

#test the code in the command line
cmd = [ 'python', 'scrabble.py', 'PENGU*?' ]
out=bytes.decode(subprocess.Popen( cmd, stdout=subprocess.PIPE ).communicate()[0])

tot_time=time.time()-start
print('Total time was {} seconds'.format(tot_time))
assert_less(tot_time, 300)
```

Output:
```
Total time was 1.4572815895080566 seconds
```

<!-- grade_id: cell-157fccf420d13535, points: 10, grade: true, solution: true -->
```python
"""Implement extra credit call run the code here
   If this cell isnt filled out - we'll assume the extra credit wasn't done
   Please write a comment on how to use your extra credit syntax also!
"""
# YOUR EXTRA CREDIT CODE HERE
```

## If you have feedback for this homework, please submit it using the link below:

http://goo.gl/forms/74yCiQTf6k